// Welcome to cpplings!
//
// This first exercise already compiles and runs. Its only job is to teach you
// the loop:
//
//   1. Open this file.
//   2. Read the code and the TODO comments.
//   3. When the exercise is solved, delete the marker comment near the bottom
//      (the line that starts with two slashes and shouts that you are not done).
//   4. Save the file. cpplings rechecks automatically.
//
// This exercise has no bug. Just remove the marker line to move on.
//
// I AM NOT DONE

#include <iostream>

int main() {
    std::cout << "Welcome to cpplings!\n";
    std::cout << "Fix the exercises, one at a time.\n";
    return 0;
}
