// The modulo operator % gives the remainder of integer division.
//
// TODO: Set `rem` to the remainder of 17 divided by 5 using the % operator.
// The test expects 2.
//
// I AM NOT DONE

#include "cpplings_test.hpp"

CPPLINGS_TEST(remainder_of_17_by_5) {
    int rem = 0;  // TODO: compute 17 % 5
    CPPLINGS_ASSERT_EQ(rem, 2);
}
